# Revision Problem Set
## Homework

**HW0.** Write a function `palindrome(s)`to check if the string `s` is a [Palindrome](https://en.wikipedia.org/wiki/Palindrome). Use iteration for this solution.


```python
def palindrome(s):
    flag = False
    s.replace(" ", "")
    for idx in range(int(len(s)/2)):
        if s[idx] != s[-(idx+1)]:
            return False
    return True
```

<!-- grade_id: hw0, points: 1, grade: true, solution: false -->
```python
assert not palindrome("moon") 
assert palindrome("noon") 
assert palindrome("a a") 
assert palindrome("ada") 
assert not palindrome("ad a") 
```

**HW1.** Create a function to implement **Bubble Sort Algorithm version 2**. You can check the pseudocode in [Bubble Sort and Insertion Sort](https://colab.research.google.com/drive/1vV2lkHiLcX-bQ6OFvciBQ_o1_hz2p4Yw). 

The function should ...

- **modify/mutate** an existing array into a sorted one
- **return** the number of comparisons made

Hint: To count the number of comparisons made you can simply increment an integer variable (counter) right before the `if` statement.

```python
#  pseudocode
#   1. n = length of array
#   2. swapped = False
#   3. As long as swapped is True, do:
#       3.1 swapped = False
#       3.2 For inner_index from 1 to n-1, do:
#           3.2.1 first_number = array[inner_index - 1]
#           3.2.2 second_number = array[inner_index]
#           3.2.3 if first_number > second_number, do:
#               3.2.3.1 swap(first_number, second_number)
#               3.2.3.2 swapped = True
```

```python
# Bubble Sort version 2

def bubble_sort(array):
    count = 0
    n = len(array)
    swapped = True
    while swapped:
        swapped = False
        for inner_idx in range(1, n):
#             print(array)
            if array[inner_idx - 1] > array[inner_idx]:
                array[inner_idx - 1], array[inner_idx] = array[inner_idx], array[inner_idx - 1]
                swapped = True
            count += 1   
    return count
```

```python

```

<!-- grade_id: hw11, points: 1, grade: true, solution: false -->
```python
array = [10,3,8,47,1,0,-39,8,4,7,6,-5]
count = bubble_sort(array)
print(array, count)
assert array == [-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
assert count == 121
```

Output:
```
[-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47] 121
```

<!-- grade_id: hw12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**HW2.** The solution in HW1 can be improved!

The n-th pass places the n-th largest element into its final place (i.e. after 1 cycle, the (1st) largest element will be in its correct position).

Hence, we can reduce subsequent iterations by only considering the other (unsorted) `n-1` elements in the array. 

Implement the algorithm of **Bubble Sort version 4** under [Bubble Sort and Insertion Sort](https://colab.research.google.com/drive/1vV2lkHiLcX-bQ6OFvciBQ_o1_hz2p4Yw). 


```python

```

```python
def bubble_sort(array):
    count = 0
    n = len(array)
    swapped = True
    while swapped:
        swapped = False
        new_n = 0
#         print('---')
        for inner_idx in range(1, n):
#             print(array)
            if array[inner_idx - 1] > array[inner_idx]:
                array[inner_idx - 1], array[inner_idx] = array[inner_idx], array[inner_idx - 1]
                swapped = True
                new_n = inner_idx
            count += 1
        n = new_n
    return count
```

```python

```

```python

```

<!-- grade_id: hw21, points: 1, grade: true, solution: false -->
```python
array = [10,3,8,47,1,0,-39,8,4,7,6,-5]
count = bubble_sort(array)
print(array, count)
assert array == [-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
assert count == 66
```

Output:
```
[-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47] 66
```

<!-- grade_id: hw22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```