# Revision Problem Set

## Cohort Problems

**CS0.** Implement factorial problem using iteration. The function should takes in an Integer input and returns and Integer output which is the factorial of the input. Recall that:

$n! = n \times (n-1) \times (n-2) \times \ldots \times 2 \times 1$

You should consider the case when $n$ is zero and one as well. 

```python
# iteration
def factorial_iteration(n):
    result = None

    if n == 0 or n == 1: 
        return 1
    else :
        result = n * factorial_iteration(n-1)
        
    return result
```

```python
factorial_iteration(11)
```

Output:
```
39916800
```

<!-- grade_id: cs01, points: 1, grade: true, solution: false -->
```python
assert factorial_iteration(0) == 1
assert factorial_iteration(1) == 1
assert factorial_iteration(5) == 120
assert factorial_iteration(7) == 5040
assert factorial_iteration(11) == 39916800
```

<!-- grade_id: cs02, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```


**CS1.** Create a function that implements **Bubble Sort version 1** (from your Notes) to sort an array of integers. The function should sort the input array in place. Refer to [Bubble Sort and Insertion Sort](https://colab.research.google.com/drive/1vV2lkHiLcX-bQ6OFvciBQ_o1_hz2p4Yw) for the pseudocode. 

```python
def myGreaterThanSwap(a, b):
    if a > b:
        return b, a
    else:
        return a, b
    

def bubble_sort(array):
    i = 0
    while i < len(array) -1:
        for index in range(len(array)-1):
            array[index], array[index+1] = myGreaterThanSwap(array[index], array[index+1])
        i += 1
    return array
```

```python

```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
array = [10,3,8,47,1,0,-39,8,4,7,6,-5]
bubble_sort(array)
print(array)
assert array == [-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
```

Output:
```
[-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** Modify CS1, so that it returns  the number of comparisons that are performed.

Hint: To count the number of comparisons, you can create an integer variable which you *increment* just before the `if` statement. 

```python
def myGreaterThanSwap(a, b, count):
    if a > b:
        return b, a, count+1
    else:
        return a, b, count+1
    
def bubble_sort(array):
    count = 0
    i = 0
    while i < len(array)-1:
        for index in range(len(array)-1):
            array[index], array[index+1], count = myGreaterThanSwap(array[index], array[index+1], count)
        i += 1
    return count
```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
array = [10,3,8,47,1,0,-39,8,4,7,6,-5]
count = bubble_sort(array)
print(array, count)
assert array == [-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47]
assert count == 121
```

Output:
```
[-39, -5, 0, 1, 3, 4, 6, 7, 8, 8, 10, 47] 121
```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3.** Create a function that implements **Insertion Sort** to sort an array of floats. The function should sort the input array in place. Refer to [Bubble Sort and Insertion Sort](https://colab.research.google.com/drive/1vV2lkHiLcX-bQ6OFvciBQ_o1_hz2p4Yw) for the pseudocode.

```python

def insertion_sort(array):
    n = len(array)
    for i in range(1, n):
        j = i
#         print('---')
        while j > 0 and array[j] < array[j-1]:
            array[j - 1], array[j] = array[j], array[j - 1]
#             print(array)
            j -= 1
```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
array = [10.3,3.8,8.4,47.1,1.0,0,-39.8,8.4,4.7,7.6,-6.5,-5.0]
insertion_sort(array)
print(array)
assert array == [-39.8, -6.5, -5.0, 0.0, 1.0, 3.8, 4.7, 7.6, 8.4, 8.4, 10.3, 47.1]
```

Output:
```
[-39.8, -6.5, -5.0, 0, 1.0, 3.8, 4.7, 7.6, 8.4, 8.4, 10.3, 47.1]
```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS4.** Write a function `gen_random_int` that generates a list of integers from `0` to `n - 1` with its sequence randomly shuffled. The function should take in an integer `n`, denoting the number of integers to be generated. 

Hint: 
- You can use `random.shuffle(mylist)`  to shuffle the elements of a list called `mylist`.
- Refer to [Python's Random module](https://docs.python.org/3.7/library/random.html) for more details on how to use `shuffle()`.

```python
import random

def gen_random_int(number, seed):
    result = None
    
    result = list(range(number))
    
    random.seed(seed)
    random.shuffle(result)
    
    return result
```

```python
output = gen_random_int(10, 100)
print(output)
assert output == [4, 0, 5, 9, 3, 1, 6, 8, 7, 2]
```

Output:
```
[4, 0, 5, 9, 3, 1, 6, 8, 7, 2]
```

```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```