# Complexity Problem Set

## Cohort Sessions

**CS1.** *Binary Heap:* Write the following functions:



- `parent_of(index)`: returns the index of node's parent
- `left_of(index)`: returns the index of node's left child
- `right_of(index)`: returns the index of node's right child
- `max_child(array, index, heap_size)`: returns the index of node's largest child. You can assume that the node has at least one child.

Hint: 
- `index` starts from 0. 
- You can refer to the pseudocode(s)??? in [Binary Heap and Heapsort](https://colab.research.google.com/drive/1D3qW_LJYwbg9e7ZlGos-eDBM0A_ftm4C).
- When finding the index of the largest child, consider the following cases:
    1. when the node only has one child
    2. when the node has two children


```python
def parent_of(index):
    return (index-1) // 2
```

<!-- grade_id: cs11, points: 1, grade: true, solution: false -->
```python
assert parent_of(1) == 0 
assert parent_of(2) == 0
assert parent_of(5) == 2
assert parent_of(6) == 2
```

<!-- grade_id: cs12, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
def left_of(index):
    return (index * 2) + 1
```

<!-- grade_id: cs13, points: 1, grade: true, solution: false -->
```python
assert left_of(0) == 1
assert left_of(1) == 3
assert left_of(3) == 7
assert left_of(6) == 13
```

<!-- grade_id: cs14, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
def right_of(index):
    return (index + 1) * 2
```

<!-- grade_id: cs15, points: 1, grade: true, solution: false -->
```python
assert right_of(0) == 2
assert right_of(1) == 4
assert right_of(3) == 8
assert right_of(5) == 12
```

<!-- grade_id: cs16, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
def max_child(array, index, heap_size):
    if right_of(index) >= heap_size:
        return left_of(index)
    else:
        if array[left_of(index)] > array[right_of(index)]:
            return left_of(index)
        else:
            return right_of(index)
```

<!-- grade_id: cs17, points: 1, grade: true, solution: false -->
```python
maxheap = [16, 14, 10, 8, 7, 9, 3, 2, 4, 1] #This is not actually a max heap cause of 2 4
assert max_child(maxheap, 0, len(maxheap)) == 1
assert max_child(maxheap, 2, len(maxheap)) == 5
assert max_child(maxheap, 3, len(maxheap)) == 8
assert max_child(maxheap, 1, len(maxheap)) == 3
```

<!-- grade_id: cs18, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS2.** *Binary Heap:* Write two functions 

- `max_heapify(array, index, size)`: that moves the node down so as to satisfy the heap property.  The first argument is the array that contains the heap. The second argument is an integer index where to start the process of heapifying. The third argument is the size of the heap in the array. This argument will be useful in heapsort algorithm where we take out the elements in the array from the heap. 

Hint: You should make use of `size` argument to determine the last element of the heap in the array rather than `len(array)`.

- `build_max_heap(heap)`: that builds the max heap from any array. This function should make use of `max_heapify()` in its definition.

Hint: You can refer to the pseudocode in [Binary Heap and Heapsort](https://colab.research.google.com/drive/1D3qW_LJYwbg9e7ZlGos-eDBM0A_ftm4C) for the above functions.

```python
# def max-heapify(A, i):
# version: 1
# Input: 
#   - A = array storing the heap
#   - i = index of the current node to restore max-heap property
# Output: None, restore the element in place
# Steps:
# 1. current_i = i # current index starting from input i
# 2. As long as ( left(current_i) < length of array), do:
#     2.1 max_child_i = get the index of largest child of the node current_i
#     2.2 if array[max_child_i] > array[current_i], do:
#         2.2.1 swap( array[max_child_i], array[current_i])
#     2.3 current_i = max_child_i # move to the index of the largest child
```

```python
def max_heapify(array, index, size):
    current_index = index
#     while left_of(current_index) < len(array): # wrong code
    while left_of(current_index) < size:  #correct code
        max_child_index = max_child(array, current_index, size)
        if array[max_child_index] > array[current_index] :
            array[max_child_index] , array[current_index] =  array[current_index], array[max_child_index]
        current_index = max_child_index

```

<!-- grade_id: cs21, points: 1, grade: true, solution: false -->
```python
result = [16, 4, 10, 14, 7, 9, 3, 2, 8, 1]
max_heapify(result, 1, len(result))
print(result)
assert result == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
```

Output:
```
[16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
```

<!-- grade_id: cs22, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

```python
# def build-max-heap(array):
# Input:
#   - array: arbitrary array of integers
# Output: None, sort the element in place
# Steps:
# 1. n = length of array
# 2. starting_index = integer(n / 2) - 1 # start from the middle or non-leaf node
# 3. For current_index in Range(from starting_index down to 0), do:
#     3.1 call max-heapify(array, current_index)
```

```python
def build_max_heap(array):
    n = len(array)
    starting_index = int(n/2) -1
#     print(f"[DEBUG] starting index : {starting_index}")
    for current_index in range(starting_index, -1, -1):
#         print(f"[DEBUG] current_index : {current_index}")
        max_heapify(array,current_index, len(array))

```

<!-- grade_id: cs23, points: 1, grade: true, solution: false -->
```python
array = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]

build_max_heap(array)
print(array)
assert array == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
```

Output:
```
[16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
```

<!-- grade_id: cs24, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS3.** *Heapsort:* Implement heapsort algorithm following the pseudocode in [Binary Heap and Heapsort](https://colab.research.google.com/drive/1D3qW_LJYwbg9e7ZlGos-eDBM0A_ftm4C).



```python
# def heapsort(array):
# Input:
#   - array: any arbitrary array
# Output: None
# Steps:
# 1. call build-max-heap(array)
# 2. heap_end_pos = length of array - 1 # index of the last element in the heap
# 3. As long as (heap_end_pos > 0), do:
#     3.1 swap( array[0], array[heap_end_pos])
#     3.2 heap_end_pos = heap_end_pos -1 # reduce heap size
#     3.3 call max-heapify(array[from index 0 to heap_end_pos inclusive], 0)
```

```python
# followed pseudocode

def heapsort(array):
    heap_end_pos = len(array) - 1
    build_max_heap(array)
    while heap_end_pos > 0:
#         print("[DEBUG] array    ",array)
        array[0], array[heap_end_pos] = array[heap_end_pos], array[0]
        heap_end_pos = heap_end_pos - 1
        max_heapify(array, 0, heap_end_pos+1)  #correct code
#         max_heapify(array[:heap_end_pos+1], 0, heap_end_pos+1) #wrong code followed the misleading pseudocode
```

```python
# Kenny's solution

# def heapsort(array):
#     heap_size = len(array)
#     build_max_heap(array)
#     while heap_size > 0:
#         print("[DEBUG] array    ",array)
#         array[0], array[heap_size-1] = array[heap_size-1], array[0]
#         heap_size = heap_size - 1
#         max_heapify(array, 0, heap_size)
```

<!-- grade_id: cs31, points: 1, grade: true, solution: false -->
```python
array = [16, 14, 10, 8, 7, 9, 3, 2, 4, 1]
heapsort(array)
print(array)
assert array == [1, 2, 3, 4, 7, 8, 9, 10, 14, 16]
```

Output:
```
[1, 2, 3, 4, 7, 8, 9, 10, 14, 16]
```

<!-- grade_id: cs32, points: 1, grade: true, solution: false -->
```python
###
### AUTOGRADER TEST - DO NOT REMOVE
###

```

**CS4.** Measure computational time of Python's built-in sort function by filling the template below. Hint: 
- You will need the function `gen_random_int()` from Week 01 Problem Set.
- Use `sorted(list)` function of Python's list See [Python's Sorting HOW TO Documentation](https://docs.python.org/3.3/howto/sorting.html). 

```python
import time
import random

def run_function(f, x):
    start = time.time()
    f(x)
    end = time.time()
    return end-start

def gen_random_int(number, seed):
    result = None
    
    result = [range(number)]
    
    random.seed(seed)
    random.shuffle(result)
    
    return result

 
time_builtin = []
# set the maximum power for 10^power number of inputs
maxpower = 7

for n in range(1, maxpower + 1):
    # create array for 10^1, 10^2, etc 
    # use seed 100
    array = None
    
    # call run_function with sorted() and array as arguments
    # result = run_function(None, None) 
    
    array = gen_random_int(100,100)
    result = run_function(heapsort, array)

    
    time_builtin.append(result)


print(time_builtin)
```

Output:
```
[4.0531158447265625e-06, 9.5367431640625e-07, 9.5367431640625e-07, 7.152557373046875e-07, 7.152557373046875e-07, 7.152557373046875e-07, 7.152557373046875e-07]
```

```python

```